# MISA
Alohan'ny mamerina dia avereno atao Run ny notebook iray manontolo. Ny fanaovana azy dia redémarrena mihitsy ny kernel aloha (jereo menubar, safidio **Kernel$\rightarrow$Restart Kernel and Run All Cells**).

Izay misy hoe `YOUR CODE HERE` na `YOUR ANSWER HERE` ihany no fenoina. Afaka manampy cells vaovao raha ilaina. Aza adino ny mameno references eo ambany raha ilaina.

## References
Eto ilay references rehetra no apetraka

---

```python
from random import randrange
import numpy as np
from sklearn.metrics import mean_squared_error, log_loss
from sklearn.datasets import load_breast_cancer, load_diabetes


def grad_check_sparse(f, x, analytic_grad, num_checks=10, h=1e-5, error=1e-9):
    """
    sample a few random elements and only return numerical
    in this dimensions
    """

    for i in range(num_checks):
        ix = tuple([randrange(m) for m in x.shape])

        oldval = x[ix]
        x[ix] = oldval + h  # increment by h
        fxph = f(x)  # evaluate f(x + h)
        x[ix] = oldval - h  # increment by h
        fxmh = f(x)  # evaluate f(x - h)
        x[ix] = oldval  # reset

        grad_numerical = (fxph - fxmh) / (2 * h)
        grad_analytic = analytic_grad[ix]
        rel_error = abs(grad_numerical - grad_analytic) / (
            abs(grad_numerical) + abs(grad_analytic)
        )
        print(
            "numerical: %f analytic: %f, relative error: %e"
            % (grad_numerical, grad_analytic, rel_error)
        )
        assert rel_error < error

def rel_error(x, y):
    """ returns relative error """
    return np.max(np.abs(x - y) / (np.maximum(1e-8, np.abs(x) + np.abs(y))))
```

# Linear regression

```python
data = load_diabetes()
X_train1, y_train1 = data.data, data.target
w1 = np.random.randn(X_train1.shape[1]) * 0.0001
b1 = np.random.randn(1) * 0.0001
```

<!-- grade_id: cell-5ffdf99ad7cdfd69, points: 0, grade: false, solution: true -->
```python
def mse_loss_naive(w, b, X, y, alpha=0):
    """
    MSE loss function WITH FOR LOOPs
    
    Returns a tuple of:
    - loss 
    - gradient with respect to weights w
    - gradient with respect to bias b
    """
    loss = 0.0
    dw = np.zeros_like(w)
    db = 0.0
    
    N, D = X.shape
    
    for i in range(N):
        diff = w @ X[i] + b - y[i]
        loss += diff**2
        db += diff
        for k in range(D):
            dw[k] += X[i][k] * diff
    
    loss *= 1/N
    db *= 2/N
    dw *= 2/N
    
    loss += alpha * (w @ w) + alpha * b * b
    dw += 2 * alpha * w
    db += 2 * alpha * b
    
    return loss, dw, np.array(db).reshape(1,)
```

## Naive Linear regression loss

<!-- grade_id: cell-079e42b153b67a60, points: 1, grade: true, solution: false -->
```python
loss, dw1, db1 = mse_loss_naive(w1, b1, X_train1, y_train1, alpha=0)

sk_loss = mean_squared_error(X_train1 @ w1 + b1, y_train1)
print("Loss error : ",rel_error(loss, sk_loss))
assert rel_error(loss, sk_loss) < 1e-9

print("Gradient check w")
# Check with numerical gradient w
f = lambda w1: mse_loss_naive(w1, b1, X_train1, y_train1, alpha=0)[0]
grad_numerical = grad_check_sparse(f, w1, dw1, 15,  error=1e-5)

print("Gradient check bias")
# Check with numerical gradient b
f2 = lambda b1: mse_loss_naive(w1, b1, X_train1, y_train1, alpha=0)[0]
grad_numerical = grad_check_sparse(f2, b1, db1, 15,  error=1e-5)
```

Output:
```
Loss error :  0.0
Gradient check w
numerical: 2.892060 analytic: 2.892060, relative error: 6.119902e-08
numerical: -3.153316 analytic: -3.153317, relative error: 1.823942e-07
numerical: -1.275044 analytic: -1.275044, relative error: 1.564797e-07
numerical: -1.376394 analytic: -1.376394, relative error: 3.654873e-08
numerical: -2.801913 analytic: -2.801913, relative error: 1.033762e-08
numerical: -0.315455 analytic: -0.315454, relative error: 2.748622e-06
numerical: -3.153316 analytic: -3.153317, relative error: 1.823942e-07
numerical: -4.145419 analytic: -4.145418, relative error: 1.176107e-07
numerical: -4.296084 analytic: -4.296087, relative error: 3.573580e-07
numerical: -3.153316 analytic: -3.153317, relative error: 1.823942e-07
numerical: 2.892060 analytic: 2.892060, relative error: 6.119902e-08
numerical: -3.234107 analytic: -3.234110, relative error: 3.335644e-07
numerical: -0.315455 analytic: -0.315454, relative error: 2.748622e-06
numerical: 2.892060 analytic: 2.892060, relati
```

## Naive Ridge regression loss

<!-- grade_id: cell-14a64cb1cfe70b98, points: 1, grade: true, solution: false -->
```python
loss, dw1, db1 = mse_loss_naive(w1, b1, X_train1, y_train1, alpha=1)

print("Gradient check w")
# Check with numerical gradient w
f = lambda w1: mse_loss_naive(w1, b1, X_train1, y_train1, alpha=1)[0]
grad_numerical = grad_check_sparse(f, w1, dw1, 15,  error=1e-5)

print("Gradient check bias")
# Check with numerical gradient b
f2 = lambda b1: mse_loss_naive(w1, b1, X_train1, y_train1, alpha=1)[0]
grad_numerical = grad_check_sparse(f2, b1, db1, 15,  error=1e-5)
```

Output:
```
Gradient check w
numerical: -1.376652 analytic: -1.376652, relative error: 4.014714e-09
numerical: 2.892002 analytic: 2.892002, relative error: 6.733758e-08
numerical: -4.145517 analytic: -4.145516, relative error: 1.196969e-07
numerical: -2.801744 analytic: -2.801744, relative error: 5.079956e-09
numerical: -4.295962 analytic: -4.295965, relative error: 3.646344e-07
numerical: -0.315212 analytic: -0.315211, relative error: 2.576344e-06
numerical: 2.892002 analytic: 2.892002, relative error: 6.733758e-08
numerical: -3.234052 analytic: -3.234054, relative error: 3.487931e-07
numerical: -4.295962 analytic: -4.295965, relative error: 3.646344e-07
numerical: -0.315212 analytic: -0.315211, relative error: 2.576344e-06
numerical: -2.801744 analytic: -2.801744, relative error: 5.079956e-09
numerical: -4.295962 analytic: -4.295965, relative error: 3.646344e-07
numerical: -4.295962 analytic: -4.295965, relative error: 3.646344e-07
numerical: -4.145517 analytic: -4.145516, relative error: 1.1969
```

<!-- grade_id: cell-1528a28f467d90c7, points: 0, grade: false, solution: true -->
```python
def mse_loss_vectorized(w, b, X, y, alpha=0):
    """
    MSE loss function WITHOUT FOR LOOPs
    
    Returns a tuple of:
    - loss 
    - gradient with respect to weights w
    - gradient with respect to bias b
    """
    
    diff = X @ w + b - y
    loss = diff @ diff
    db = diff.sum()
    dw = X.T @ diff
    
    N, _ = X.shape
    
    loss *= 1/N
    db *= 2/N
    dw *= 2/N
    
    loss += alpha * (w @ w) + alpha * b * b
    dw += 2 * alpha * w
    db += 2 * alpha * b
    
    return loss, dw, np.array(db).reshape(1,)
```

## Vectorised Linear regression loss

<!-- grade_id: cell-41637ca21c8f938d, points: 1, grade: true, solution: false -->
```python
loss, dw1, db1 = mse_loss_vectorized(w1, b1, X_train1, y_train1, alpha=0)

sk_loss = mean_squared_error(X_train1 @ w1 + b1, y_train1)
print("Loss error : ",rel_error(loss, sk_loss))
assert rel_error(loss, sk_loss) < 1e-9

print("Gradient check w")
# Check with numerical gradient w
f = lambda w1: mse_loss_vectorized(w1, b1, X_train1, y_train1, alpha=0)[0]
grad_numerical = grad_check_sparse(f, w1, dw1, 15,  error=1e-5)

print("Gradient check bias")
# Check with numerical gradient b
f2 = lambda b1: mse_loss_vectorized(w1, b1, X_train1, y_train1, alpha=0)[0]
grad_numerical = grad_check_sparse(f2, b1, db1, 15,  error=1e-5)
```

Output:
```
Loss error :  6.256300376732055e-17
Gradient check w
numerical: -4.145418 analytic: -4.145418, relative error: 7.911924e-09
numerical: -3.234110 analytic: -3.234110, relative error: 3.899087e-09
numerical: -2.801913 analytic: -2.801913, relative error: 4.279740e-08
numerical: -1.275044 analytic: -1.275044, relative error: 1.381878e-08
numerical: -2.801913 analytic: -2.801913, relative error: 4.279740e-08
numerical: 2.892060 analytic: 2.892060, relative error: 3.314495e-08
numerical: -1.376395 analytic: -1.376394, relative error: 9.560738e-08
numerical: -1.376395 analytic: -1.376394, relative error: 9.560738e-08
numerical: -1.275044 analytic: -1.275044, relative error: 1.381878e-08
numerical: -0.315454 analytic: -0.315454, relative error: 1.538098e-07
numerical: -3.153317 analytic: -3.153317, relative error: 1.950318e-08
numerical: -1.275044 analytic: -1.275044, relative error: 1.381878e-08
numerical: -2.801913 analytic: -2.801913, relative error: 4.279740e-08
numerical: -1.275044 analy
```

## Vectorized ridge regression loss

<!-- grade_id: cell-bb0ff99c19a85e0f, points: 1, grade: true, solution: false -->
```python
loss, dw1, db1 = mse_loss_vectorized(w1, b1, X_train1, y_train1, alpha=1)

print("Gradient check w")
# Check with numerical gradient w
f = lambda w1: mse_loss_vectorized(w1, b1, X_train1, y_train1, alpha=1)[0]
grad_numerical = grad_check_sparse(f, w1, dw1, 15,  error=1e-5)

print("Gradient check bias")
# Check with numerical gradient b
f2 = lambda b1: mse_loss_vectorized(w1, b1, X_train1, y_train1, alpha=1)[0]
grad_numerical = grad_check_sparse(f2, b1, db1, 15,  error=1e-5)
```

Output:
```
Gradient check w
numerical: -1.274979 analytic: -1.274979, relative error: 2.648522e-08
numerical: -1.376652 analytic: -1.376652, relative error: 1.281167e-07
numerical: -1.274979 analytic: -1.274979, relative error: 2.648522e-08
numerical: -4.295965 analytic: -4.295965, relative error: 2.589974e-08
numerical: -4.145516 analytic: -4.145516, relative error: 1.000067e-08
numerical: -1.274979 analytic: -1.274979, relative error: 2.648522e-08
numerical: 2.892002 analytic: 2.892002, relative error: 2.700828e-08
numerical: -0.315211 analytic: -0.315211, relative error: 2.046841e-08
numerical: -4.145516 analytic: -4.145516, relative error: 1.000067e-08
numerical: -1.376652 analytic: -1.376652, relative error: 1.281167e-07
numerical: -0.315211 analytic: -0.315211, relative error: 2.046841e-08
numerical: -1.376652 analytic: -1.376652, relative error: 1.281167e-07
numerical: 2.892002 analytic: 2.892002, relative error: 2.700828e-08
numerical: -4.295965 analytic: -4.295965, relative error: 2.5899
```

# Logistic regression

```python
def sigmoid(z):
    return 1 / (1 + np.exp(-z))

data = load_breast_cancer()
X_train2, y_train2 = data.data, data.target
w2 = np.random.randn(X_train2.shape[1]) * 0.0001
b2 = np.random.randn(1) * 0.0001
```

# Naive

<!-- grade_id: cell-6ec156568b0c6e29, points: 0, grade: false, solution: true -->
```python
def log_loss_naive(w, b, X, y, alpha=0):
    """
    log loss function WITH FOR LOOPs
    
    Returns a tuple of:
    - loss 
    - gradient with respect to weights w
    """
    loss = 0.0
    dw = np.zeros_like(w)
    db = 0.0
    
    N, D = X.shape
    
    for i in range(N):
        z = w @ X[i] + b
        
        p = sigmoid(z)
        loss -= y[i] * np.log(p) + (1 - y[i]) * np.log(1 - p) 
        
        c = p - y[i]
        db += c
        for k in range(D):
            dw[k] += c * X[i][k]
            
    loss *= 1/N
    dw *= 1/N
    db *= 1/N
        
    loss += alpha * (w @ w) + alpha * b * b
    dw += 2 * alpha * w
    db += 2 * alpha * b
    
    return loss, dw, np.array(db).reshape(1,)
```

<!-- grade_id: cell-d3c078eb2449ee61, points: 1, grade: true, solution: false -->
```python
y_pred_0 = sigmoid(X_train2 @ w2 + b2)
y_pred = np.vstack([1-y_pred_0, y_pred_0]).T
sk_loss = log_loss(y_train2, y_pred)

loss, dw2, db2 = log_loss_naive(w2, b2, X_train2, y_train2, alpha=0)
print("Loss error : ",rel_error(loss, sk_loss))
assert rel_error(loss, sk_loss) < 1e-9

print("Gradient check w")
# Check with numerical gradient w
f = lambda w2: log_loss_naive(w2, b2, X_train2, y_train2, alpha=0)[0]
grad_numerical = grad_check_sparse(f, w2, dw2, 15, error=1e-4)

print("Gradient check bias")
# Check with numerical gradient b
f2 = lambda b2: log_loss_naive(w2, b2, X_train2, y_train2, alpha=0)[0]
grad_numerical = grad_check_sparse(f2, b2, db2, 15,  error=1e-5)
```

Output:
```
Loss error :  2.3130982781671467e-16
Gradient check w
numerical: -0.001494 analytic: -0.001494, relative error: 1.514755e-08
numerical: 0.389991 analytic: 0.389991, relative error: 6.354444e-11
numerical: -0.004245 analytic: -0.004245, relative error: 2.486878e-09
numerical: 0.018737 analytic: 0.018737, relative error: 1.286459e-09
numerical: -0.445248 analytic: -0.445248, relative error: 4.835900e-09
numerical: 0.018737 analytic: 0.018737, relative error: 1.286459e-09
numerical: 6.241341 analytic: 6.241343, relative error: 9.427985e-08
numerical: 0.008765 analytic: 0.008765, relative error: 1.782394e-09
numerical: -0.000609 analytic: -0.000609, relative error: 5.445799e-08
numerical: -0.003045 analytic: -0.003045, relative error: 1.043201e-09
numerical: -0.000609 analytic: -0.000609, relative error: 5.445799e-08
numerical: 0.018737 analytic: 0.018737, relative error: 1.286459e-09
numerical: 158.236769 analytic: 158.238570, relative error: 5.691834e-06
numerical: 0.361658 analytic: 0.3
```

# Naive with regulariztion

<!-- grade_id: cell-b91805308ad91ec5, points: 1, grade: true, solution: false -->
```python
loss, dw2, db2 = log_loss_naive(w2, b2, X_train2, y_train2, alpha=1)

print("Gradient check w")
# Check with numerical gradient w
f = lambda w2: log_loss_naive(w2, b2, X_train2, y_train2, alpha=1)[0]
grad_numerical = grad_check_sparse(f, w2, dw2, 15, error=1e-4)

print("Gradient check bias")
# Check with numerical gradient b
f2 = lambda b2: log_loss_naive(w2, b2, X_train2, y_train2, alpha=1)[0]
grad_numerical = grad_check_sparse(f2, b2, db2, 15,  error=1e-5)
```

Output:
```
Gradient check w
numerical: -0.382712 analytic: -0.382712, relative error: 1.350333e-08
numerical: 0.361933 analytic: 0.361933, relative error: 3.054748e-09
numerical: -0.000271 analytic: -0.000271, relative error: 4.098044e-08
numerical: 0.000315 analytic: 0.000315, relative error: 4.081562e-08
numerical: -0.004392 analytic: -0.004392, relative error: 5.665868e-10
numerical: -0.088934 analytic: -0.088934, relative error: 1.128813e-10
numerical: 0.012018 analytic: 0.012018, relative error: 1.580740e-09
numerical: -0.007629 analytic: -0.007629, relative error: 2.512169e-09
numerical: 0.054000 analytic: 0.054000, relative error: 3.804675e-11
numerical: 10.408111 analytic: 10.408112, relative error: 4.990051e-08
numerical: 0.390267 analytic: 0.390267, relative error: 6.366544e-11
numerical: 0.054000 analytic: 0.054000, relative error: 3.804675e-11
numerical: 0.361933 analytic: 0.361933, relative error: 3.054748e-09
numerical: 0.018669 analytic: 0.018669, relative error: 1.365690e-09
numer
```

# Vectorized

<!-- grade_id: cell-a96e9a6d51919ffc, points: 0, grade: false, solution: true -->
```python
def log_loss_vectorized(w, b,X, y, alpha=0):
    """
    log loss function WITHOUT FOR LOOPs
    
    Returns a tuple of:
    - loss 
    - gradient with respect to weights w
    """
    
    N, D = X.shape
    
    Z = X @ w + b
    P = sigmoid(Z)
    D = P - y
    
    loss = - (Z @ y) - (np.log(1 - P)).sum()
    dw = X.T @ D
    db = D.sum()
    
    loss *= 1/N
    dw *= 1/N
    db *= 1/N
        
    loss += alpha * (w @ w) + alpha * b * b
    dw += 2 * alpha * w
    db += 2 * alpha * b
    
    return loss, dw, np.array(db).reshape(1,)
```

<!-- grade_id: cell-ca14e49ada130789, points: 1, grade: true, solution: false -->
```python
y_pred_0 = sigmoid(X_train2 @ w2 + b2)
y_pred = np.vstack([1-y_pred_0, y_pred_0]).T
sk_loss = log_loss(y_train2, y_pred)

loss, dw2, db2 = log_loss_vectorized(w2, b2, X_train2, y_train2, alpha=0)
print("Loss error : ",rel_error(loss, sk_loss))
assert rel_error(loss, sk_loss) < 1e-9

print("Gradient check w")
# Check with numerical gradient w
f = lambda w2: log_loss_vectorized(w2, b2, X_train2, y_train2, alpha=0)[0]
grad_numerical = grad_check_sparse(f, w2, dw2, 15, error=1e-4)

print("Gradient check bias")
# Check with numerical gradient b
f2 = lambda b2: log_loss_vectorized(w2, b2, X_train2, y_train2, alpha=0)[0]
grad_numerical = grad_check_sparse(f2, b2, db2, 15,  error=1e-5)
```

Output:
```
Loss error :  7.710327593890492e-17
Gradient check w
numerical: -0.382700 analytic: -0.382700, relative error: 1.347838e-08
numerical: 85.791315 analytic: 85.791846, relative error: 3.099678e-06
numerical: 0.361658 analytic: 0.361658, relative error: 3.094193e-09
numerical: 0.001655 analytic: 0.001655, relative error: 3.904197e-09
numerical: -0.382700 analytic: -0.382700, relative error: 1.347838e-08
numerical: 0.001655 analytic: 0.001655, relative error: 3.904197e-09
numerical: 0.001655 analytic: 0.001655, relative error: 3.904197e-09
numerical: 3.023797 analytic: 3.023798, relative error: 1.066256e-07
numerical: -0.000609 analytic: -0.000609, relative error: 4.346670e-09
numerical: 0.000460 analytic: 0.000460, relative error: 3.363482e-09
numerical: 3.023797 analytic: 3.023798, relative error: 1.066256e-07
numerical: -0.003045 analytic: -0.003045, relative error: 7.796916e-10
numerical: -0.004245 analytic: -0.004245, relative error: 7.825290e-10
numerical: -0.004515 analytic: -0.0045
```

# Vectorized with regularization

<!-- grade_id: cell-bce082fccf8b7057, points: 1, grade: true, solution: false -->
```python
loss, dw2, db2 = log_loss_vectorized(w2, b2, X_train2, y_train2, alpha=1)

print("Gradient check w")
# Check with numerical gradient w
f = lambda w2: log_loss_vectorized(w2, b2, X_train2, y_train2, alpha=1)[0]
grad_numerical = grad_check_sparse(f, w2, dw2, 15, error=1e-4)

print("Gradient check bias")
# Check with numerical gradient b
f2 = lambda b2: log_loss_vectorized(w2, b2, X_train2, y_train2, alpha=1)[0]
grad_numerical = grad_check_sparse(f2, b2, db2, 15,  error=1e-5)
```

Output:
```
Gradient check w
numerical: 0.000315 analytic: 0.000315, relative error: 3.200546e-09
numerical: 0.054000 analytic: 0.054000, relative error: 3.804688e-11
numerical: -0.004571 analytic: -0.004571, relative error: 7.271969e-10
numerical: 10.408111 analytic: 10.408112, relative error: 4.990184e-08
numerical: -0.382712 analytic: -0.382712, relative error: 1.348157e-08
numerical: 0.000315 analytic: 0.000315, relative error: 3.200546e-09
numerical: 6.241310 analytic: 6.241311, relative error: 9.427906e-08
numerical: 0.018669 analytic: 0.018669, relative error: 2.761526e-11
numerical: -0.000271 analytic: -0.000271, relative error: 6.932716e-11
numerical: 0.022426 analytic: 0.022426, relative error: 1.207473e-10
numerical: 0.824292 analytic: 0.824292, relative error: 2.381495e-09
numerical: -0.445147 analytic: -0.445147, relative error: 4.816253e-09
numerical: -0.445147 analytic: -0.445147, relative error: 4.816253e-09
numerical: -0.007629 analytic: -0.007629, relative error: 3.440512e-11
num
```

# Gradient descent for Linear models

<!-- grade_id: cell-485e52c0efd4f4a9, points: 0, grade: false, solution: true -->
```python
class LinearModel():
    def __init__(self):
        self.w = None
        self.b = None

    def train(self, X, y, learning_rate=1e-3, alpha=0, num_iters=100, batch_size=200, verbose=False):
        N, d = X.shape
        
        if self.w is None: # Initialization
            self.w = 0.001 * np.random.randn(d)
            self.b = 0.0

        # Run stochastic gradient descent to optimize w
        
        loss_history = []
        for it in range(num_iters):
            
            index = np.random.choice(N, batch_size, replace=False)
            
            X_batch = X[index, :]
            y_batch = y[index]
                                                        
            # evaluate loss and gradient
            loss, dw, db = self.loss(X_batch, y_batch, alpha)
            loss_history.append(loss)

            self.w -= learning_rate * dw
            self.b -= learning_rate * db
            
            if verbose and it % 10000 == 0:
                print("iteration %d / %d: loss %f" % (it, num_iters, loss))
                
        return loss_history

    def predict(self, X):
        pass

    def loss(self, X_batch, y_batch, reg):
        pass

class LinearRegressor(LinearModel):
    """ Linear regression """

    def loss(self, X_batch, y_batch, alpha):
        return mse_loss_vectorized(self.w, self.b, X_batch, y_batch, alpha)
    
    def predict(self, X):
        return X @ self.w + self.b

class LogisticRegressor(LinearModel):
    """ Linear regression """

    def loss(self, X_batch, y_batch, alpha):
        return log_loss_vectorized(self.w, self.b, X_batch, y_batch, alpha)
    
    def predict(self, X):
        """ Return prediction labels vector of 0 or 1 """
        return X @ self.w + self.b > 0
```

## Linear regression with gradient descent

<!-- grade_id: cell-92f36a3b387a4277, points: 2, grade: true, solution: false -->
```python
from sklearn.linear_model import LinearRegression

sk_model = LinearRegression(fit_intercept=True)
sk_model.fit(X_train1, y_train1)
sk_pred = sk_model.predict(X_train1)
sk_mse = mean_squared_error(sk_pred, y_train1)

model = LinearRegressor()
model.train(X_train1, y_train1, num_iters=75000, batch_size=64, learning_rate=1e-2, verbose=True)
pred = model.predict(X_train1)
mse = mean_squared_error(pred, y_train1)

print("MSE scikit-learn:", sk_mse)
print("MSE gradient descent model :", mse)
assert mse - sk_mse < 100
```

Output:
```
iteration 0 / 75000: loss 26914.469491
iteration 10000 / 75000: loss 3925.371916
iteration 20000 / 75000: loss 2979.772292
iteration 30000 / 75000: loss 3109.754402
iteration 40000 / 75000: loss 2934.168923
iteration 50000 / 75000: loss 2907.236718
iteration 60000 / 75000: loss 2879.823378
iteration 70000 / 75000: loss 3476.886247
MSE scikit-learn: 2859.69634758675
MSE gradient descent model : 2884.6180303248834
```

## Logistc regression with gradient descent

<!-- grade_id: cell-925a2ddb5c1ba7ff, points: 2, grade: true, solution: false -->
```python
from sklearn.linear_model import LogisticRegression
from sklearn.preprocessing import StandardScaler

scaler = StandardScaler()
X_train2 = scaler.fit_transform(X_train2)

sk_model = LogisticRegression(fit_intercept=True)
sk_model.fit(X_train2, y_train2)
sk_pred = sk_model.predict(X_train2)
sk_log_loss = log_loss(sk_pred, y_train2)

model = LogisticRegressor()
model.train(X_train2, y_train2, num_iters=75000, batch_size=64, learning_rate=1e-3, verbose=True)
pred = model.predict(X_train2)
model_log_loss = log_loss(pred, y_train2)

print("Log-loss scikit-learn:", sk_log_loss)
print("Log-loss gradiet descent model :", model_log_loss)
print("Error :", rel_error(sk_log_loss, model_log_loss))
assert rel_error(sk_log_loss, model_log_loss) < 1e-7
```

Output:
```
iteration 0 / 75000: loss 0.693513
iteration 10000 / 75000: loss 0.107225
iteration 20000 / 75000: loss 0.044944
iteration 30000 / 75000: loss 0.071879
iteration 40000 / 75000: loss 0.148246
iteration 50000 / 75000: loss 0.033898
iteration 60000 / 75000: loss 0.205370
iteration 70000 / 75000: loss 0.072091
Log-loss scikit-learn: 0.44341928598210933
Log-loss gradiet descent model : 0.44341928598210933
Error : 0.0
```

```python

```